This exercise is mostly just to ensure that everything is setup correctly, and to get familiarized to the assignment submission process - there is minimal programming in this assignment.

```python
# You must run this cell before starting your assignment

!pip install -q otter-grader

import otter
grader = otter.Notebook("assignment1.ipynb")

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

Output:
```
[1m[[0m[34;49mnotice[0m[1;39;49m][0m[39;49m A new release of pip is available: [0m[31;49m24.3.1[0m[39;49m -> [0m[32;49m25.0[0m
[1m[[0m[34;49mnotice[0m[1;39;49m][0m[39;49m To update, run: [0m[32;49mpython3 -m pip install --upgrade pip[0m
```

**Question 1:** 
Set a variable called `var` to 5 (as an integer).\
Set a variable called `name` to your name (as a str).\
Set a variable called `age` to your age (as an integer).

```python
# Code goes here:
var = int(5)
name = "Walker Turley"
age = 20
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

You have completed all the questions in the assignment, you can now download this file and upload it to Gradescope.

1) To download this, click on this file on the Explorer pane on the left, and click on Download. 
2) Open this link on a new tab: here: https://www.gradescope.com/courses/934147/assignments/5674111.
3) Upload the file you downloaded where it says to in the Gradescope assignment. 

Please feel free to ask me any questions about the assignment or the submision process by raising your hand!



```python

```

Tests for q1 (points: 3)
```
>>> isinstance(var, int) == True
True
```
